Please read the following instructions thoroughly. Neglecting to do so may result in missed points.

### Preamble
**Reminder**: Homeworks are due by 11:59PM ET on Sundays.

Before you turn this problem set in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

### Naming conventions
Be sure the filename of your notebook is in the following form:

    <uni>_<assignment>_<details [optional]>.<extension>
    
For example:

    lr3086_hw01.ipynb
    lr3086_hw01_complete.ipynb
    LR3086_HW01.ipynb
    
To rename a notebook, in the menubar, select File$\rightarrow$Rename. The extension for notebook files, `.ipynb`, will already be appended to the filename, but will be hidden from view within the notebook.
    
This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `lr3086_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

### What Format To Submit In

Most homeworks are in Jupyter notebooks. Once you've finished your homework, unless specified otherwise, please download your work as an `.ipynb` file to your local machine, then upload it to Courseworks when complete (in the menubar, select File$\rightarrow$Download as$\rightarrow$Notebook).

**Failure to submit a Jupyter notebook will result in a grade of zero for the assignment.**

### Grading

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors.

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem.

Empty un-editable cells in an assignment are there for a reason. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

### Getting Help

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Ed Discussions or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Ed Discussions and not via email. Answers given via Ed Discussions will help you as well as other students. Thus, emails will always have a lower priority for response than Ed Discussions questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in Courseworks faster.

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

---

# Homework 2: Fundamentals Continued

Total questions: 11<br/>
Total points: 14

Some notes about this homework:

* You should not need any imports or additional libraries to implement the solutions this homework, except for **question 9**. However, you may import other modules or libraries if you wish to write any test code for your solutions.
* Some problems have cells containing some code with `### BEGIN TESTS` and `### END TESTS`. These are you help you test your solutions. Execute the cells, and if they raise an error, your solution is not correct. Do not delete these cells. Note that we will run additional tests against your solutions for correctness.

## Question 1

Define a function called `underline` which takes two parameters: a string, and then a character to use to underline the string with. The function should `return` a string with 2 lines denoted by newline characters, where the first line is the string provided and the second is the underline character repeated the same length as the string itself.

Example input to `underline`:

```py
my_string = "I love autumn!"
ch = "+"
```

Expected returned string from calling `underline`:

```py
'I love autumn!\n++++++++++++++\n'
```

There will be no credit awarded if you use the `print` function.


[1 point]

<!-- grade_id: cell-fa141663d0ed07d7, points: 0, grade: false, solution: true -->
```python
def underline(string, ch):
    # YOUR CODE HERE
    underline=ch * len(string)
    s=f"{string}\n{underline}\n"
    return s
    #raise NotImplementedError()
    
#underline("ciao", "+")
```

<!-- grade_id: cell-b301d957b48b2da1, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
result = underline("I love autumn!", "+")

assert result == "I love autumn!\n++++++++++++++\n"
### END TESTS
```

<!-- grade_id: cell-d7e8fbd840494efa, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-f01bb0932efc46f6, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

## Question 2

Implement the function `evens` that returns all the even numbers between `0` and a given number `x`. Zero should be included in the result but `x` should not be included unless `x` is zero. You may return a `list` or any other container type. `x` may be negative, in which case an empty `list` must be returned.

[1 point]

<!-- grade_id: cell-7d8a3589ac5f0ac4, points: 0, grade: false, solution: true -->
```python
def evens(x):
    # YOUR CODE HERE
    my_list=list()
    if x<0:
        return my_list
    elif x==0:
        my_list.append(0)
        return my_list
    else:
        for t in range(x):
            if t%2==0:
                my_list.append(t)
        return my_list
    
#number=evens(3)
#number
    #raise NotImplementedError()
```

<!-- grade_id: cell-0ec921ce986e1340, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-3ed6b952ba18b807, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-5abd8faea8cd8101, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-34be4caf495be8ac, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

## Question 3
Write a function `compress` that takes a string of characters, counts each adjacent occurrence of a character, and produces a compressed string with each character and its count. For example:

`'aaaabbbbcc'` would compress to `'a4b4c2'`.

You can assume that the source string will be lowercase and that only consecutively adjacent characters should be compressed -- in otherwords, `aabbaa` should compress to `a2b2a2`, not `a4b2`.

[2 points]

<!-- grade_id: cell-184a8b7b502b6efa, points: 0, grade: false, solution: true -->
```python


def compress(letter_string):
    char=""
    string=""
    counter=""
    for t in range(len(letter_string)):
        if(char!=letter_string[t]):
            string=f"{string}{char}{counter}"
            char=letter_string[t]
            counter=1
        else:
            counter=counter+1
        
    return f"{string}{char}{counter}"       
            
    # YOUR CODE HERE
#compress('')
    
    #raise NotImplementedError()
```

<!-- grade_id: cell-f12c5fbbb8ae6490, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
assert compress('aaaabbbbcccc') == 'a4b4c4'
### END TESTS
```

<!-- grade_id: cell-1639e9c5b3229082, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
assert compress('') == ''
### END TESTS
```

<!-- grade_id: cell-e545067e23350627, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
assert compress('cccddee') == 'c3d2e2'
### END TESTS
```

<!-- grade_id: cell-bad44bfc6357c312, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-9aa4c63d1d4c41cb, points: 1, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

## Question 4
Write a function `endpoints` that takes a list of numbers (eg. `[5, 10, 15, 20, 25]`) and returns a new `list` of only the first and last elements of the given list (eg. `[5, 25]`). If the input list is `[5]`, the returned should be `[5,5]`. The function should return an empty list if an empty list is passed in.

[1 point]

<!-- grade_id: cell-4b0beddf6eecb07e, points: 0, grade: false, solution: true -->
```python
def endpoints(numbers):
    # YOUR CODE HERE
    result=list()
    if numbers==result:
        return result
    length=len(numbers)
    if length==1:
        return [numbers[0],numbers[0]]
    else:
        return [numbers[0], numbers[-1]]
    #raise NotImplementedError()
    
#r=endpoints([5,10,15,35,25])
#r
```

<!-- grade_id: cell-e336488834b9dcc5, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
assert endpoints([5, 10, 15, 20, 25]) == [5, 25]
### END TESTS
```

<!-- grade_id: cell-bed663010f6a7325, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
assert endpoints([5]) == [5, 5]
### END TESTS
```

<!-- grade_id: cell-ab0ea0b635f7b663, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
assert endpoints([]) == []
### END TESTS
```

<!-- grade_id: cell-2b6ecbf12b92a452, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

## Question 5
Write a function `common` that consumes two lists of numbers, and that produces a list that contains only the common elements (without duplicates). Make sure your program works on two lists of different sizes and returns the common elements in sorted ascending order.

```py
a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89]
b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]

# result
[1, 2, 3, 5, 8, 13]
```

[1 point]

<!-- grade_id: cell-e0a0bee134bf6c41, points: 0, grade: false, solution: true -->
```python
import collections

def common(a, b):
    # YOUR CODE HERE
    numb_count = collections.Counter(a)
    result=list()
    for number in numb_count:
        if number in b:
            result.append(number)
    result.sort()
    return result
     
#a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89]
#b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]

#h=common(b, a)
#h

   # raise NotImplementedError()
```

<!-- grade_id: cell-ff1a9c45bd92afcd, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
a = [1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89]
b = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13]

assert [1, 2, 3, 5, 8, 13] == common(a, b)
### END TESTS
```

<!-- grade_id: cell-3a9684fc4029bb39, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-11a30cf3a405184f, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-9d48a711fc5c0aac, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-dd36efae1fe39ac4, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

## Question 6

Define a function named `merge` that takes in two dictionaries and merges them into one dictionary. If the two dictionaries happen to have common keys, then the first dictionary should take precedent. The returned dictionary doesn't matter for this question.

```py
dict1 = {"a": 10, "b": 8}
dict2 = {"c": 6, "d": 4}
# result
{"a": 10, "b": 8, "c": 6, "d": 4}

dict3 = {"a": 10, "b": 8}
dict4 = {"b": 12, "c": 8}
# result
{"a": 10, "b": 8, "c": 8}
```

[1 point]

<!-- grade_id: cell-ace481ddf07aeed8, points: 0, grade: false, solution: true -->
```python
def merge(dict1, dict2):
    # YOUR CODE HERE
    result=dict1
    for element in dict2:
        if element not in dict1:
            result[element]=dict2.get(element)
    
    return result
    

   # raise NotImplementedError()
```

<!-- grade_id: cell-d5dbb312cf5e6a7b, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
dict1 = {"a": 10, "b": 8}
dict2 = {"c": 6, "d": 4}
expected = {"a": 10, "b": 8, "c": 6, "d": 4}
actual = merge(dict1, dict2)
assert expected == actual
### END TESTS
```

<!-- grade_id: cell-20939ceef4565313, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
dict3 = {"a": 10, "b": 8}
dict4 = {"b": 12, "c": 8}
expected = {"b": 8, "a": 10, "c": 8}
actual = merge(dict3, dict4)
assert expected == actual
### END TESTS
```

<!-- grade_id: cell-89744c8383a81495, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

## Question 7

We have covered some fundamental string manipulation operations in class, such as `str.strip` and `str.replace`. There are [many many more](https://docs.python.org/3/library/stdtypes.html#string-methods) which you should familiarize yourself with.

Because of the richness of these methods, parsing a string into meaningful subcomponents is often a breeze.

Imagine a (multiline) string like:

```py
"""
Columbia to Launch 25 Million Dollar AI-based Climate Modeling Center
---------------------------------------------------------------------
Author: Kim Martineau

Funded by the National Science Foundation, the center will leverage big data and machine learning to improve climate projections.

It should also motivate societies to invest in policies and infrastructure to confront rising seas and warmer temperatures.
"""
```
 
Implement the `parse` function so that it parses the above news article snippet into:

```py
{
    "title": "Columbia to Launch 25 Million Dollar AI-based Climate Modeling Center",
    "author": "Kim Martineau",
    "body": "Funded by the National Science Foundation, the center will leverage big data and machine learning to improve climate projections.\n\nIt should also motivate societies to invest in policies and infrastructure to confront rising seas and warmer temperatures.",
}
```
i.e. a dict where the title of the article and its body have been extracted into two separate values.

Articles may be missing (i.e. not have) an author. As another example:

```py
"""
Columbia Pledges That All Future Campus Construction Will Be Fossil Free
------------------------------------------------------------------------

As Climate Week NYC begins, the university explores creating a fully electrified campus. The Columbia Climate School is university partner of the weeklong climate showcase.
"""
```

should be parsed into:

```
{
    "title": "Columbia Pledges That All Future Campus Construction Will Be Fossil Free",
    "author": None,
    "body": "As Climate Week NYC begins, the university explores creating a fully electrified campus. The Columbia Climate School is university partner of the weeklong climate showcase.",
}
```
    
You may assume that:

* the title is always followed by a line which underlines it
* the next line will begin with exactly the contents "Author: " followed by the author's name, if the article has an author
* any further lines are the body of the article
* the body of the article may contain multiple lines
* leading and trailing whitespace around the body can be ignored

[3 points]

<!-- grade_id: cell-7efd8eb3a945d729, points: 0, grade: false, solution: true -->
```python
def parse(article):
    # YOUR CODE HERE
    art=article.strip()
    titleEnd=art.find('\n-')
    result=dict()
    title=""
    for x in range(titleEnd):
        title=f"{title}{art[x]}"
    
    char="-"
    result["title"]=title
    art=art.replace(f"{title}\n{char*len(title)}\n", "")
    vettore=art.split("\n",1)
    author=vettore[0]
    if author.find("Author:")==0:
        author=author.replace("Author: ","")
        result["author"]=author
    else:
        result["author"]=None
    
    body=vettore[1].strip()
        
    result["body"]=body
    return result

    
    
  #  raise NotImplementedError()
```

<!-- grade_id: cell-1cc5f8295f1646b2, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
article = """
Columbia to Launch 25 Million Dollar AI-based Climate Modeling Center
---------------------------------------------------------------------
Author: Kim Martineau

Funded by the National Science Foundation, the center will leverage big data and machine learning to improve climate projections.

It should also motivate societies to invest in policies and infrastructure to confront rising seas and warmer temperatures.
"""

result = parse(article)
result["body"] = result["body"].strip()

assert result == {
    "title": "Columbia to Launch 25 Million Dollar AI-based Climate Modeling Center",
    "author": "Kim Martineau",
    "body": "Funded by the National Science Foundation, the center will leverage big data and machine learning to improve climate projections.\n\nIt should also motivate societies to invest in policies and infrastructure to confront rising seas and warmer temperatures.",
}
### END TESTS
```

<!-- grade_id: cell-34bdfb37b05edb9c, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
article = """
Columbia Pledges That All Future Campus Construction Will Be Fossil Free
------------------------------------------------------------------------

As Climate Week NYC begins, the university explores creating a fully electrified campus. The Columbia Climate School is university partner of the weeklong climate showcase.
"""

result = parse(article)
result["body"] = result["body"].strip()

assert result == {
    "title": "Columbia Pledges That All Future Campus Construction Will Be Fossil Free",
    "author": None,
    "body": "As Climate Week NYC begins, the university explores creating a fully electrified campus. The Columbia Climate School is university partner of the weeklong climate showcase.",
}
### END TESTS
```

<!-- grade_id: cell-5588dd72c0904f3b, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-c8e99b7b1964c5c6, points: 1, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

## Question 8
Python's standard library has a module called `string` which contains a number of useful string constants which can be tedious to manually type out.

For example, [string.ascii_uppercase](https://docs.python.org/3/library/string.html#string.ascii_uppercase) is a pre-built string containing the Latin (i.e. English) alphabet in uppercase -- so the letters 'ABC...XYZ'.

In the cell below, write a single line of code that imports the string module such that someone can use `ascii_uppercase` without needing to prefix it by string. -- i.e. so that simply `ascii_uppercase` is the above referenced string.
    
[0.5 Points]

<!-- grade_id: cell-b32b21ca0c5ebcba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from string import ascii_uppercase

#raise NotImplementedError()

```

<!-- grade_id: cell-81bf828af634cd65, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

## Question 9
A businessman named Alan has an office calendar with 12 free meeting slots, numbered simply from `1` through `12`.

Alan has 3 stakeholders, Liam, Mia and Sofia that he needs to meet with on an ongoing basis, each of whom put meetings on Alan's calendar. No meeting can overlap with any other.

Alan needs a way of deciding whether his calendar is fully booked for a given day.

Write a function `is_alan_fully_booked`, which takes as input 3 separate sets which represent which slots each of Liam, Mia and Sofia have booked on his calendar, and which returns `True` if Alan has no free slots on his calendar, and `False` if he has a remaining free slot.

As an example -- `is_alan_fully_booked(liam={1, 7, 8}, mia={2, 3, 4, 5, 6}, sofia={9, 10, 11, 12})` should return `True`, and `is_alan_fully_booked(liam={1, 2, 3, 4, 5, 6, 7}, mia=set(), sofia=set())` should return `False`.

[1.5 point]

<!-- grade_id: cell-97a85e467ea1f00e, points: 0, grade: false, solution: true -->
```python
def is_alan_fully_booked(liam, mia, sofia):
    # YOUR CODE HERE
    union=liam | mia | sofia
    if union=={1,2,3,4,5,6,7,8,9,10,11,12}:
        return True
    else:
        return False

    #raise NotImplementedError()
```

<!-- grade_id: cell-fe146e44bd228523, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
liam = {1, 7, 8}
mia = {2, 3, 4, 5, 6}
sofia = {9, 10, 11, 12}
assert is_alan_fully_booked(liam, mia, sofia) == True
### END TESTS
```

<!-- grade_id: cell-f9373b9f89a46c6c, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
liam = {1, 2, 3, 4, 5, 6, 7}
mia = set()
sofia = set()
assert is_alan_fully_booked(liam, mia, sofia) == False
### END TESTS
```

<!-- grade_id: cell-8d74a09271993ab6, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

## Question 10

In addition to knowing whether he is fully booked, Alan needs to know what slots he has remaining in a given day.

Define a function `free_slots` which again takes each of the 3 stakeholders' booked slots and returns a set containing all the slots which are not booked by any of the 3 stakeholders.

As an example -- `free_slots(liam={1, 7, 8}, mia=set(), sofia={9, 10, 11, 12})` should return `{2, 3, 4, 5, 6}`, and `free_slots(liam={1, 2, 3, 4, 5, 6, 7, 12}, mia=set(), sofia=set())` should return `{8, 9, 10, 11}`.

[1 point]

<!-- grade_id: cell-d3837974e2ec00bd, points: 0, grade: false, solution: true -->
```python
def free_slots(liam, mia, sofia):
    # YOUR CODE HERE
    union=liam | mia | sofia
    result={1,2,3,4,5,6,7,8,9,10,11,12}-union
    return result


    #raise NotImplementedError()
```

<!-- grade_id: cell-90d815492ba13e4c, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
liam = {1, 7, 8}
mia = set()
sofia = {9, 10, 11, 12}
expected_free_slots = {2, 3, 4, 5, 6}

assert free_slots(liam, mia, sofia) == expected_free_slots
### END TESTS
```

<!-- grade_id: cell-f45f96eeb0c069b2, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
liam = {1, 2, 3, 4, 5, 6, 7, 12}
mia = set()
sofia = set()
expected_free_slots = {8, 9, 10, 11}

assert free_slots(liam, mia, sofia) == expected_free_slots
### END TESTS
```

<!-- grade_id: cell-8459fcef488fee7a, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

## Question 11
Finally, Alan also wants to know who out of Liam, Mia and Sofia are taking most of his time on a given day.

Write a function `booked_most_meetings` which calculates who has booked the most of Alan's time on a given day.

If multiple people have booked an equal amount of Alan's time, return any arbitrary single person's name.

As an example, `booked_most_meetings(liam={1, 7, 8}, mia=set(), sofia={9, 10, 11, 12})` should return `"sofia"` who has booked 4.

`booked_most_meetings(liam={1}, mia={2}, sofia={3})` may return any of `"liam"`, `"mia"`, or `"sofia"`, all of whom have booked an equal amount of Alan's time.

[1 point]

<!-- grade_id: cell-fe56b467015d30a5, points: 0, grade: false, solution: true -->
```python
def booked_most_meetings(liam, mia, sofia):
    # YOUR CODE HERE
    my_map={
        "liam":len(liam),
        "mia":len(mia),
        "sofia":len(sofia)
    }
    max=0
    result=""
    for element in my_map:
        if my_map[element]>max:
            max=my_map[element]
            result=element
    
    return element

    #raise NotImplementedError()
```

<!-- grade_id: cell-14b336688ccfc7d9, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
liam = {1, 7, 8}
mia = set()
sofia = {9, 10, 11, 12}

assert booked_most_meetings(liam, mia, sofia) == "sofia"
### END TESTS
```

<!-- grade_id: cell-018d5a7ace76b8a0, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
# DO NOT DELETE OR OVERWRITE
liam = {1}
mia = {2}
sofia = {3}

expected_result = {"liam", "mia", "sofia"}
# Since it can be any of them since it's all equal, we're just ensuring the result 
# is one of the three names
assert booked_most_meetings(liam, mia, sofia) in expected_result
### END TESTS
```

<!-- grade_id: cell-770f0a87ab829a29, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-5bba4f41f8118b18, points: 0.25, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY FOR HIDDEN TESTS - DO NOT DELETE OR OVERWRITE
```