Please read the following instructions thoroughly. Neglecting to do so may result in missed points.

### Preamble
**Reminder**: Homeworks are due by 11:59PM ET on Sundays.

Before you turn this problem set in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

### Naming conventions
Be sure the filename of your notebook is in the following form:

    <uni>_<assignment>_<details [optional]>.<extension>
    
For example:

    lr3086_hw01.ipynb
    lr3086_hw01_complete.ipynb
    LR3086_HW01.ipynb
    
To rename a notebook, in the menubar, select File$\rightarrow$Rename. The extension for notebook files, `.ipynb`, will already be appended to the filename, but will be hidden from view within the notebook.
    
This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `lr3086_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

### What Format To Submit In

Most homeworks are in Jupyter notebooks. Once you've finished your homework, unless specified otherwise, please download your work as an `.ipynb` file to your local machine, then upload it to Courseworks when complete (in the menubar, select File$\rightarrow$Download as$\rightarrow$Notebook).

**Failure to submit a Jupyter notebook will result in a grade of zero for the assignment.**

### Grading

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors.

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem.

Empty un-editable cells in an assignment are there for a reason. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

### Getting Help

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Ed Discussions or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Ed Discussions and not via email. Answers given via Ed Discussions will help you as well as other students. Thus, emails will always have a lower priority for response than Ed Discussions questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in Courseworks faster.

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

---

# Homework 3: List Comprehensions, Exceptions, File I/O, and Recursion

Total questions: 6 <br/>
Total points: 10

Some questions have imports already for you. Additional imports of modules or libraries **are not allowed** in your solutions. However, you may import other modules or libraries if you wish to write any test code for your solutions.

Remember: Unless otherwise stated in the question, be sure that your final answer to a question includes a `return` from your implemented function. Do not print. You can add extra notebook cells to test your solution, and print the results. We just need your functions to `return`. Refer to the past lecture notes for an example of what we need.

## Question 1
Write a function called `scalar_multiply` which takes 2 arguments -- an `int` and a sequence (e.g. a `list`) of `int`s, of any length.

The function should return a new `list` where each element in the sequence has been multiplied by the `int`.

The function should raise a `ValueError` if the sequence given is empty, with any message provided.

The function should otherwise always return a `list`, even if the provided sequence is not a list.

Do not use a `for`-loop or a `while`-loop, you must use a list comprehension. **No points will be awarded if `for`-/`while`-loops are used.**

As an example, `scalar_multiply(3, [2, 4, 5])` should return the list` [6, 12, 15]`, where each element has been multiplied by `3`.

Calling `scalar_multiply(11, [])` should raise a `ValueError`.

[2 points]

<!-- grade_id: cell-4fbbd3615eb93ed9, points: 0, grade: false, solution: true -->
```python
def scalar_multiply(scalar, sequence):
    # YOUR CODE HERE
    if len(sequence)!=0:
        result=[scalar*element for element in sequence]
        return result
    else:
        raise ValueError("There is an error, check it")
    
            
   # raise NotImplementedError()
    
#result=scalar_multiply(3,{2,4,5})
#print(result)
```

<!-- grade_id: cell-558efc9f599b7c27, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert scalar_multiply(3, [2, 4, 5]) == [6, 12, 15]
### END TESTS
```

<!-- grade_id: cell-c32d61a912c8b281, points: 1.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY - DO NOT DELETE OR OVERWRITE
```

## Question 2

_Preamble: In class, we worked with text files and learned how to read and write files. Be sure to review the lecture notes. When working with numerical data, you will often encounter data in the form of [CSV files](https://en.wikipedia.org/wiki/Comma-separated_values). Take a moment to familiarize yourself with the [`csv` module](https://docs.python.org/3/library/csv.html) from Python's standard library, which provides an easy way to read & write CSV-based data._

Below is a function called `download_park_data` that downloads data from New York State about State Park annual attendence since 2003, and returns it as a string. You can take a look at the data directly through their [website](https://data.ny.gov/Recreation/State-Park-Annual-Attendance-Figures-by-Facility-B/8f3n-xj78/data).

Implement a function called `get_palisades_attendance` that does not take in any arguments. This function should do the following:

1. **Get data**: Call the `download_park_data` function.
2. **Write data to file**: Save the data that the `download_park_data` function returns by writing it to a file on your computer with a context manager. Name the written file `"downloaded_data.csv"`. It should not be saved in any particular directory. That is, use `open("downloaded_data.csv", "w")` with no file path. You do not need to use the `csv` module here.
3. **Read the file**: Then, use another context manager to read the lines of the file you just saved. You should be left with a `list` of `str`s.  You do not need to use the `csv` module here.
4. **Parse the file**: After you've read the lines of the file, use the `csv` module to parse the lines of data into a single `list` containing multiple of `list`s that represents rows of data.
5. **Calculate result**: Calculate the total attendance (`total_attendance`) from 2003 to 2020 (inclusive) for the `"Appalachian Trail"` in the `"Palisades"` OPRHP region. Return the value. The total attendence should be 156,506 people.

Remember, you may **not** use additional imports for your implementation (you can import what you want in order to test your implementation though, if you wish).

[2 points]

<!-- grade_id: cell-a79e2598d389f3b6, points: 0, grade: false, solution: true -->
```python
import csv
import requests

def download_park_data():
    response = requests.get(
        "https://data.ny.gov/api/views/8f3n-xj78/rows.csv?accessType=DOWNLOAD&sorting=true"
    )
    data = response.text
    return data


def get_palisades_attendance():
    data = download_park_data()
    # YOUR CODE HERE
    with open("downloaded_data.csv", "w") as new:
        new.write(data)
    with open("downloaded_data.csv") as file:
        ok = file.readlines()
    csv_reader = csv.reader(ok)
    result=list()
    for row in csv_reader:
        result.append(row)
    attendence=0
    for row in result[1:]:
        if int(row[0]) in range(2003,2021)  and row[1]=="Palisades" and row[3]=="Appalachian Trail" :
            attendence=attendence+int(row[4])
    return attendence
   # raise NotImplementedError()


#get_palisades_attendance()
```

<!-- grade_id: cell-fd2833e7c51b92c4, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
assert 156506 == get_palisades_attendance()
### END TESTS
```

<!-- grade_id: cell-7fd3b1398d6802d7, points: 2, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY - DO NOT DELETE OR OVERWRITE
```

## Question 3

Write a recursive function `sum_list` to find the sum of all elements in a list. You can assume the list is flat. You **can not** use the built-in function `sum()`. You **can not** use iteration via a `for` or `while` loop. Your function must return an integer. You can assume you are passed a list of real integers. If your solution uses exceptions, the except statements must not be bare (i.e. the must specify the errors they can catch). You must handle the case where the given list is empty; the sum would be zero in this case.

[2 points]

<!-- grade_id: cell-7f5b01a5937961bd, points: 0, grade: false, solution: true -->
```python
def sum_list(a_list):
    # YOUR CODE HERE
    if a_list==[]:
        return 0
    else:
        return sum_list(a_list[:-1])+a_list[-1]
    


    # raise NotImplementedError()
```

<!-- grade_id: cell-263728976e123f6c, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert sum_list([]) == 0
### END TESTS
```

<!-- grade_id: cell-d21a7be16b3fd4f6, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-5e1ac4f8bee8695f, points: 1, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY - DO NOT DELETE OR OVERWRITE
```

## Question 4
Create a recursive function `fib` to find the nth [Fibonacci number](https://en.wikipedia.org/wiki/Fibonacci_number). Assume that the the sequence starts with `0, 1, 1, ...`.

```py
>>> fib(0)
0
>>> fib(5)
5
>>> fib(10)
55
```

If your solution uses exceptions, the except statements must not be bare (ie. the must specify the errors they can catch).

[2 points]

<!-- grade_id: cell-315cf28b73c195e3, points: 0, grade: false, solution: true -->
```python
def fib(n):
    # YOUR CODE HERE
    if n==0:
        return 0
    elif n==1:
        return 1
    else:
        return fib(n-1)+fib(n-2)
   # raise NotImplementedError()
```

<!-- grade_id: cell-b62e5b681aa6bf04, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
assert fib(0) == 0
assert fib(5) == 5
assert fib(10) == 55
### END TESTS
```

<!-- grade_id: cell-6f77d73604953548, points: 1, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY - DO NOT DELETE OR OVERWRITE
```

## Question 5

Write a function `verify` that takes a string and returns the string if it contains at least one number, one uppercase letter, and one lowercase letter. The function should raise a `ValueError` otherwise. The `ValueError` should have the **exact** message `"Must contain a number, an uppercase letter, and a lowercase letter"`.

[1 point]

<!-- grade_id: cell-ebadbf8ae994dfa6, points: 0, grade: false, solution: true -->
```python
def verify(string):
    # YOUR CODE HERE
    number=False
    upper=False
    lower=False
    for c in string:
        if c.islower():
            lower=True
        if c.isupper():
            upper=True
        if c.isnumeric():
            number=True
    if number and upper and lower:
        return string
    else:
        raise ValueError("Must contain a number, an uppercase letter, and a lowercase letter")
        

    
    #raise NotImplementedError()
```

<!-- grade_id: cell-4965c28d46a8fe5a, points: 1, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY - DO NOT DELETE OR OVERWRITE
```

## Question 6

Implement a function called `is_ten` that takes in one argument, an `int`, and returns `True` if the argument is equal to `10`, and `False` if not.

Then, implement a second function called `test_is_ten` that takes in no arguments, and calls the first function, `is_ten` at least twice. For each time you call `is_ten`, use `assert` statements to test the expected return value equals actual return value.

[1 point]

<!-- grade_id: cell-044f254c957716f7, points: 0, grade: false, solution: true -->
```python
def is_ten(x):
    # YOUR CODE HERE
    #raise NotImplementedError()
    if x==10:
        return True
    else:
        return False
    
def test_is_ten():
    # YOUR CODE HERE
    for x in range(5,11):
        if x!=10:
            assert is_ten(x)==False
        else:
            assert is_ten(x)==True
   # raise NotImplementedError()

#test_is_ten()
```

<!-- grade_id: cell-744a5046c8b96cbd, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY - DO NOT DELETE OR OVERWRITE
```

<!-- grade_id: cell-814f19a6082746c4, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY EMPTY - DO NOT DELETE OR OVERWRITE
```