Please read the following instructions thoroughly. Neglecting to do so may result in missed points.

### Preamble
**Do Not Talk During Quizzes.** Do not talk at all once a quiz has begun. You may talk again once you leave the room after the quiz. The one exception to this rule is if you need to ask the exam proctor a question. Any talking during a quiz qualifies the student for disciplinary action.

You have 30 minutes to complete the quiz.

Before you turn this problem set in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

### Permissions for this quiz

* No collaborating. Your work must be your own.
* You may use the internet and Jupyter Notebooks from lectures, homeworks, and notes.
* Upload the solved notebook as you would a homework **before the Courseworks timer is up**. Completed exams will not be accepted!

### Completing this quiz

To begin the quiz, **rename the notebook** to be of the form `<uni>_quiz1`. For example, mine would be `lr3086_quiz1`. **If your files are not named with this format, you should expect a grade of zero for the assignment.**

When you are done with your exam:

1. Save this exam.
2. Download this exam as an `.ipynb` file.
3. Upload the `.ipynb` file to Courseworks **before the timer is up**. Completed exams will not be accepted!

Please note, there are several cells in this Jupyter notebook that are hidden and read only. Do not attempt to remove them or edit them. They are used in grading your notebook.

* DO remove the "Not Implemented" lines when implementing your solution.
* DO test all cells to make sure they run successfully, and that they terminate in 30 seconds or less. To do so:
    * if using Jupyter Lab, go to "Run -> Restart Kernel and Run All Cells...", OR
    * if using Jupyter Notebook (not Jupyter Lab) go to "Kernel -> Restart & Clear Output" then "Cell -> Run All".

Please reach out to a CA or the professor in class for any clarifying questions.


---

# Homework 4: Object Oriented Programming

Total Questions: 6<br/>
Total Points: 10

## Context

Questions 1 through 5 will have you practice some object oriented programming by implementing an **Address Book** for managing contact information.

You will be asked to implement a number of new types, and some methods on them for implementing address book functionality you likely are already familiar with.

When asked to modify an existing type in an exercise, you should completely redefine the type -- in other words, if you answered:

```py
class Something:
    def __init__(self):
        pass
```

for the first question, and are then asked to add a new `foo` attribute to `Something` instances in question 2, your answer to question 2 should be:

```py
class Something:
    def __init__(self):
        self.foo = 12
```

where you have completely redefined the Something class with the newly needed attribute.

## Question 1

We first will start with a type which represents each **contact** we store in the address book.

We'll want a contact to have a (full) name which is required, zero or more phone numbers, and zero or more email addresses for the contact.

Here's an example of how someone will want to use our class:

```py
karina = Contact(
    name="Karina Whitley", 
    addresses=[], 
    email_addresses=["kwhitley@example.org"]
)
```

Write a `Contact` class supporting the above, which additionally re-exposes the `name`, `addresses`, and `email_addresses` arguments as public attributes. In other words, in the above example, `karina.name` should return `"Karina Whitley"`, the string that was passed in.

Remember that both `addresses` and `email_addresses` support 0 or more addresses. To make providing 0 of these more convenient, your class should support leaving those parameters off, meaning:

```py
tayte = Contact(name="Tayte Whitaker")
```

should create a contact whose `addresses` and `email_addresses` are set to a falsy value (any falsy value is acceptable, like `None` or an empty list, `[]`).

[2 points]

<!-- grade_id: cell-58ef6d0845f50544, points: 0, grade: false, solution: true -->
```python
class Contact:
    # YOUR CODE HERE
    def __init__(self, name, addresses=[], email_addresses=[]):
        self.name=name
        self.addresses=addresses
        self.email_addresses=email_addresses
    
    @property    
    def name(self):
        return self.__name
    
    @property
    def addresses(self):
        return self.__addresses
    
    @property
    def email_addresses(self):
        return self.__email_addresses
    
    @name.setter
    def name(self, new_value):
        self.__name=new_value
        
    @addresses.setter
    def addresses(self, new_value):
        self.__addresses=new_value
    
    @email_addresses.setter
    def email_addresses(self, new_value):
        self.__email_addresses=new_value
    
   # raise NotImplementedError()
```

<!-- grade_id: cell-4086caa50583bb00, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
karina = Contact(
    name="Karina Whitley", 
    addresses=[], 
    email_addresses=["kwhitley@example.org"]
)
assert "Karina Whitley" == karina.name
assert not karina.addresses
assert ["kwhitley@example.org"] == karina.email_addresses

### END TESTS
```

<!-- grade_id: cell-cff1aad36e1aae6f, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
edwin = Contact(
    name="Edwin Alder", 
    addresses=["100 Park Avenue South\nNew York, NY 10003"], 
    email_addresses=["ealder@example.org"]
)

assert (edwin.name, edwin.addresses, edwin.email_addresses) == (
    "Edwin Alder",
    ["100 Park Avenue South\nNew York, NY 10003"],
    ["ealder@example.org"],
)
### END TESTS
```

<!-- grade_id: cell-ec15216391d25859, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
tae = Contact(name="Tae Danae Ruby")
assert tae.name == "Tae Danae Ruby"
assert not tae.addresses
assert not tae.email_addresses
### END TESTS
```

<!-- grade_id: cell-bd391629be5f0670, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
christine = Contact(
    name="Christine Robins", 
    email_addresses=["crobins@example.com", "crobins@example.edu"]
)
assert christine.name == "Christine Robins"
assert not christine.addresses
assert christine.email_addresses == ["crobins@example.com", "crobins@example.edu"]
### END TESTS
```

## Question 2

Now that we have a type which represents an individual contact (person), we want a type which represents our address book which will contain lots of contacts.

Write an `AddressBook` class which takes a list of contacts (defined as above), and can be used as such:

```py
lynn = Contact(name="Lynn Root")
book = AddressBook(contacts=[lynn])
```

Give your address book class a _method_ called `add_contact` which can add a contact to the address book even after creating one.

In other words:

```py
sai = Contact(
    name="Sai Lee", 
    email_addresses=["sai@example.com"]
)
book.add_contact(sai)
```
which should add the new contact to the address book.

[2 points]

<!-- grade_id: cell-6c24414d136a3f22, points: 0, grade: false, solution: true -->
```python
class AddressBook:
    # YOUR CODE HERE
    def __init__(self, contacts=[]):
        self.contacts=contacts
        
    @property
    def contacts(self):
        return self.__contacts
    
    @contacts.setter
    def contacts(self, new_value):
        self.__contacts=new_value
        
    def add_contact(self, new_contact):
        self.contacs=self.contacts.append(new_contact)
   # raise NotImplementedError()



```

<!-- grade_id: cell-3ac167f97aea38ca, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
lynn = Contact(name="Lynn Root")
book = AddressBook(contacts=[lynn])
book.add_contact(Contact(name="Sai Lee", email_addresses=["sai@example.com"]))
### END TESTS
```

<!-- grade_id: cell-4040a0858c2c2d70, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
book = AddressBook(contacts=[])
book.add_contact(lynn)
### END TESTS
```

## Question 3

Our address book can have contacts added to it, but we have no way of looking up their information.

Let's add one.

Redefine the `AddressBook` class such that it also has a `get_contact_by_name` method which takes a name (a `str`) and returns the contact with the given name -- i.e. it returns an instance of `Contact`.

You can assume that all names are unique, so no two contacts will have the same exact name for this exercise.

Implement `get_contact_by_name` method so that you only have to provide the string of the contact's name and not a `Contact` instance.

Be sure to also include the implementation for the `add_contact` method you defined from question 2.

[2 points]

<!-- grade_id: cell-762a12f6641e3f64, points: 0, grade: false, solution: true -->
```python
class AddressBook:
    # YOUR CODE HERE
    def __init__(self, contacts=[]):
        self.contacts=contacts
        
    @property
    def contacts(self):
        return self.__contacts
    
    @contacts.setter
    def contacts(self, new_value):
        self.__contacts=new_value
        
    def add_contact(self, new_contact):
        self.contacs=self.contacts.append(new_contact)
    
    def get_contact_by_name(self, name):
        for contact in self.contacts:
            if contact.name==name:
                return contact
        
        return None
    
        
    
    
    #raise NotImplementedError()
```

<!-- grade_id: cell-a526d1180a86e905, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
lynn = Contact(name="Lynn Root")
book = AddressBook(contacts=[lynn])
assert book.get_contact_by_name("Lynn Root") == lynn
### END TESTS
```

<!-- grade_id: cell-8a64b693f8bafd7e, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
sai = Contact(name="Sai Lee", email_addresses=["sai@example.com"])
book.add_contact(sai)
assert book.get_contact_by_name("Lynn Root") == lynn
assert book.get_contact_by_name("Sai Lee") == sai
### END TESTS
```

## Question 4

We can now both add contacts to the address book as well as look them up.

But we want to be able to signal that a contact doesn't exist if we attempt to look up one by name that is not present in the address book.

Define an exception called `NoSuchContact`, and raise this exception in `get_contact_by_name` if it is called with a name that does not match any contact. You should include the name of the (non-existing) contact in the exception message for `NoSuchContact`. 

Don't forget, you'll have to re-define the `AddressBook` class in order to contain this new logic.

[2 points]

<!-- grade_id: cell-e86904ff06b67d34, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class NoSuchContact(Exception):
    pass

class AddressBook:
    # YOUR CODE HERE
    def __init__(self, contacts=[]):
        self.contacts=contacts
        
    @property
    def contacts(self):
        return self.__contacts
    
    @contacts.setter
    def contacts(self, new_value):
        self.__contacts=new_value
        
    def add_contact(self, new_contact):
        self.contacs=self.contacts.append(new_contact)
    
    def get_contact_by_name(self, name):
        
        for contact in self.contacts:
            if contact.name==name:
                return contact
       
        
        raise NoSuchContact(f"Can't find {name}")
    
#raise NotImplementedError()
```

<!-- grade_id: cell-a98e218fe0d1fd39, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
lynn = Contact(name="Lynn Root")
book = AddressBook(contacts=[lynn])
assert book.get_contact_by_name("Lynn Root") == lynn
try:
    book.get_contact_by_name("Unknown Contact")
except NoSuchContact:
    pass
else:
    assert False, "Expected an exception!"
### END TESTS
```

<!-- grade_id: cell-a153e8f8472e06bf, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
lynn = Contact(name="Lynn Root")
book = AddressBook(contacts=[lynn])
assert book.get_contact_by_name("Lynn Root") == lynn
try:
    book.get_contact_by_name("Unknown Contact")
except NoSuchContact as e:
    assert "Unknown Contact" in str(e)
else:
    assert False, "Expected an exception!"
### END TESTS
```

<!-- grade_id: cell-5de2fe39e3024da1, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
try:
    book.get_contact_by_name("Another Unknown Contact")
except NoSuchContact as e:
    assert "Another Unknown Contact" in str(e)
else:
    assert False, "Expected an exception!"
### END TESTS
```

## Question 5

Write an implementation of the `__repr__` special method which for a `Contact` shows the name of the contact.

For a contact named `"Sara Loggins"`, we want to see `"Contact(name='Sara Loggins')"`. You should not include any other contact fields in the output. Be sure that your repr output exactly matches the example shown!

[1 point]

<!-- grade_id: cell-f319c54ef28284b0, points: 0, grade: false, solution: true -->
```python
class Contact:
    # YOUR CODE HERE
   
    def __init__(self, name, addresses=[], email_addresses=[]):
        self.name=name
        self.addresses=addresses
        self.email_addresses=email_addresses
    
    @property    
    def name(self):
        return self.__name
    
    @property
    def addresses(self):
        return self.__addresses
    
    @property
    def email_addresses(self):
        return self.__email_addresses
    
    @name.setter
    def name(self, new_value):
        self.__name=new_value
        
    @addresses.setter
    def addresses(self, new_value):
        self.__addresses=new_value
    
    @email_addresses.setter
    def email_addresses(self, new_value):
        self.__email_addresses=new_value
    
    def __repr__(self):
        return f"Contact(name='{self.name}')"
    #raise NotImplementedError()
```

<!-- grade_id: cell-8682a765f98c409d, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert repr(Contact(name="Gandalf Hernandez")) == "Contact(name='Gandalf Hernandez')"
### END TESTS
```

<!-- grade_id: cell-1950365786765918, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert repr(Contact(name="Lynn Root")) == "Contact(name='Lynn Root')"
### END TESTS
```

## Question 6

Create a class `Str` that inherits `str`. Add a method to the subclass that checks if the string does not start with a given string. Call that method `notstartswith`. This class should only take up 3 lines of code however you will not be graded on the length of your class.

[1 point]

<!-- grade_id: cell-63e0e838fde67499, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Str(str):
    def notstartswith(self,string):
        return not self.startswith(string)
        
#raise NotImplementedError()
```

<!-- grade_id: cell-437707ddadd1642a, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
what_does_the_dog_say = "woof"
word = Str(what_does_the_dog_say)

assert word.notstartswith("w") is False
### END TESTS
```

<!-- grade_id: cell-3fb8e83ec56679c5, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK
```

<!-- grade_id: cell-b59882bc0ba6b144, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK
```

```python

```