Please read the following instructions thoroughly. Neglecting to do so may result in missed points.

### Preamble
**Reminder**: Homeworks are due by 11:59PM ET on Sundays.

Before you turn this problem set in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

### Naming conventions
Be sure the filename of your notebook is in the following form:

    <uni>_<assignment>_<details [optional]>.<extension>
    
For example:

    lr3086_hw01.ipynb
    lr3086_hw01_complete.ipynb
    LR3086_HW01.ipynb
    
To rename a notebook, in the menubar, select File$\rightarrow$Rename. The extension for notebook files, `.ipynb`, will already be appended to the filename, but will be hidden from view within the notebook.
    
This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `lr3086_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

### What Format To Submit In

Most homeworks are in Jupyter notebooks. Once you've finished your homework, unless specified otherwise, please download your work as an `.ipynb` file to your local machine, then upload it to Courseworks when complete (in the menubar, select File$\rightarrow$Download as$\rightarrow$Notebook).

**Failure to submit a Jupyter notebook will result in a grade of zero for the assignment.**

### Grading

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors.

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem.

Empty un-editable cells in an assignment are there for a reason. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

### Getting Help

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Ed Discussions or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Ed Discussions and not via email. Answers given via Ed Discussions will help you as well as other students. Thus, emails will always have a lower priority for response than Ed Discussions questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in Courseworks faster.

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

---

# Homework 5: Data Structures & Algorithms

Total questions: 3<br/>
Total points: 10

There are many details in each of the questions' prompts – be sure to read closely. Run the tests below your solutions to ensure you've implemented them correctly.

## Question 1

Linked lists are a collection of nodes connect by a pointer. Each node holds a value and a pointer to the next node or to some unique value like `None` which indidcates the end of the list.

Create a class `Node` that is instantiated with a value and a reference to the next node in the linked list. The value of the node should be accessible via an attribute named `value` on the node. A reference to the next node in the list should be accessible via an attribute named `next` on the node. If the node is the last node in a linked list, the node's `next` value should be `None`. Implement a `__repr__` method that returns a human-readable string representing the node in the format of `<Node(X)>` where `X` is the value of the node instance.

Create a function `insert` that takes a reference to the first node in a linked list, a new `value`, and a position and inserts the a new node with the given value at the given position in the list. The first node of the updated list should be returned. If the position is farther than the length of the list, an `IndexError` should be raised.

Create a function `pop` that takes a reference to the first node in a linked list and a position and removes the node at that position in the linked list. This function should return two items: the popped node and the head of the linked list in a tuple. If there are no more nodes to pop, raise an `IndexError`.

Create a function `stringify_linked_list` that takes a reference to the first node of a linked list and returns a printable string of all nodes in the linked list. See the test cases below for how to format the expected string output.

[6 points]

<!-- grade_id: cell-297c228f59d53764, points: 0, grade: false, solution: true -->
```python
class Node:
    # YOUR CODE HERE
    #raise NotImplementedError()
    def __init__(self, value, reference=None):
        self.value=value
        self.next=reference
    
    @property
    def value(self):
        return self.__value
    
    @property
    def next(self):
        return self.__next
    
    @value.setter
    def value(self, value):
        self.__value=value
        
    @next.setter
    def next(self, reference):
        self.__next=reference
        
    
    
    def __repr__(self):
        return f"<Node ({self.value})>"
    
def insert(head, value, position):
    # YOUR CODE HERE
   # raise NotImplementedError()
    node=head
    if position==0:
        new=Node(value, node)
        return new
        
    i=0
    while i<position:
        if i==position-1:
            new=Node(value, node.next)
            node.next=new
            return head
        
        node=node.next
        i=i+1
        if node==None:
            raise IndexError
        


def pop(head, position):
    # YOUR CODE HERE
    #raise NotImplementedError()
    i=0
    node=head
    
    
    
    if position==0:
        new_head=head.next
        head.next=None
        return (head, new_head)
    
    
    while i<position:
        if i==position-1:
            removed=node.next
            if removed==None:
                raise IndexError
            new_next=removed.next
            removed.next=None
            node.next=new_next
            return (removed, head)
         
        i=i+1
        node=node.next
        
    

def stringify_linked_list(head):
    result=""
    node=head
    while True:
        result=f"{result}{repr(node)} -> "
        
        node=node.next
        if node==None:
            result=f"{result}None"
            return result
        
        
        
    # YOUR CODE HERE
   # raise NotImplementedError()


```

<!-- grade_id: cell-0500bff60c227ce2, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert repr(Node(-1, None)) == '<Node (-1)>'
### END TESTS
```

<!-- grade_id: cell-9ae9d5f6d9143039, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
n1 = Node(4, None)
assert n1.value == 4
assert n1.next is None
### END TESTS
```

<!-- grade_id: cell-6351a00c8f778f7a, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
n2 = Node(8, n1)
assert n2.value == 8
assert n2.next is n1
### END TESTS
```

<!-- grade_id: cell-ba6716422a08ea47, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
head = None
for i in range(3):
    head = Node(i, head)

assert stringify_linked_list(head) == '<Node (2)> -> <Node (1)> -> <Node (0)> -> None'
### END TESTS
```

<!-- grade_id: cell-2c90d3298772ac60, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
head = None
for i in range(3):
    head = Node(i, head)
head = insert(head, 0, 0)
assert stringify_linked_list(head) == '<Node (0)> -> <Node (2)> -> <Node (1)> -> <Node (0)> -> None'
### END TESTS
```

<!-- grade_id: cell-06af3d7c2fb34e01, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
head = None
for i in range(3):
    head = Node(i, head)
head = insert(head, 0, 1)
assert stringify_linked_list(head) == '<Node (2)> -> <Node (0)> -> <Node (1)> -> <Node (0)> -> None'
### END TESTS
```

<!-- grade_id: cell-398175917e71b524, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
head = None
for i in range(3):
    head = Node(i, head)
head = insert(head, -1, 3)
assert stringify_linked_list(head) == '<Node (2)> -> <Node (1)> -> <Node (0)> -> <Node (-1)> -> None'
### END TESTS
```

<!-- grade_id: cell-02ea17cc3951797f, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
head = None
for i in range(3):
    head = Node(i, head)

try:
    insert(head, 0, 4)
    assert False, 'Should have failed with an IndexError'
except IndexError:
    pass
### END TESTS
```

<!-- grade_id: cell-9678d46b04e5932a, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
head = None
for i in range(3):
    head = Node(i, head)
    
popped_node, head = pop(head, 1)
assert popped_node.value == 1, popped_node.value
assert stringify_linked_list(head) == '<Node (2)> -> <Node (0)> -> None'
### END TESTS
```

<!-- grade_id: cell-dabf452ecdc4ed30, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
head = None
for i in range(3):
    head = Node(i, head)
    
popped_node, head = pop(head, 0)
assert popped_node.value == 2, popped_node.value
assert stringify_linked_list(head) == '<Node (1)> -> <Node (0)> -> None'
### END TESTS
```

<!-- grade_id: cell-4fe6d4c54645c3f1, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
head = None
for i in range(3):
    head = Node(i, head)
    
popped_node, head = pop(head, 2)
assert popped_node.value == 0
assert stringify_linked_list(head) == '<Node (2)> -> <Node (1)> -> None'
### END TESTS
```

<!-- grade_id: cell-06b279c562c99441, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
head = None
for i in range(3):
    head = Node(i, head)

try:
    pop(head, 3)
    assert False, 'Should have failed with an IndexError'
except IndexError:
    pass
### END TESTS
```

## Question 2

Complete the implementation of bubble sort.

All sorting operations must be done on the list passed in. No new lists can be created. The function does not need to return the sorted list since the sorting will be done in-place (ie. using the same memory as the original list). You can not use the `.append` or `.extend` methods on a list. You must use the bubble sort algorithm. Assume that all members of `my_int_list` are integers.

[2 points]

<!-- grade_id: cell-492aa1fcb5a95e1a, points: 0, grade: false, solution: true -->
```python
def bubble_sort(my_int_list):
    swap_occurred = True  # kick start the while loop
    while swap_occurred:
        swap_occurred = False
        c=0
        for i in range(len(my_int_list) - 1):
            a=my_int_list[i]
            b=my_int_list[i+1]
            if a>b:
                my_int_list[i]=b
                my_int_list[i+1]=a
            
            if i>0 and c>my_int_list[i]:
                swap_occured=False
            
            c=my_int_list[i]
              
            # YOUR CODE HERE
            #raise NotImplementedError()
```

<!-- grade_id: cell-a5e25c1caa2cc61e, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a_list = [3, 1, 2]
bubble_sort(a_list)
assert a_list == [1, 2, 3]
### END TESTS
```

<!-- grade_id: cell-47bcc8c8dd347e6f, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a_list = [1, 2, 3]
bubble_sort(a_list)
assert a_list == [1, 2, 3]
### END TESTS
```

<!-- grade_id: cell-c3a47f9fb66cfcbe, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
a_list = []
bubble_sort(a_list)
assert a_list == []

a_list = [1]
bubble_sort(a_list)
assert a_list == [1]
### END TESTS
```

<!-- grade_id: cell-4b9eb986925e39c4, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

## Question 3

Complete the function `binary_search`.

Return `True` if `item` is in `sorted_list` and `False` if not. Your function must implement the `binary_search` algorithm using recursion. Assume that all members of `sorted_int_list` are integers. If you use except clauses, they must not be bare (ie. the must catch one or more specific error types).

[2 points]

<!-- grade_id: cell-45d486d546fb325f, points: 0, grade: false, solution: true -->
```python
def binary_search(item, sorted_int_list):
    # YOUR CODE HERE
    #raise NotImplementedError()
    if len(sorted_int_list)==0:
        return False
    
    start = 0 
    end = len(sorted_int_list)-1
    mid = (start + end)//2

  
    if (item == sorted_int_list[mid]):
        return True
   
    if (item > sorted_int_list[mid]):
        return binary_search(item, sorted_int_list[mid+1:] ) 
   
        return binary_search(item, sorted_int_list[:mid])
    return False
```

<!-- grade_id: cell-4b3f43f12a6039c4, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert binary_search(4, [1, 2, 3, 4, 5])
assert not binary_search(2.5, [1, 2, 3, 4, 5])
assert not binary_search(5, [])
assert binary_search(5, [5])
### END TESTS
```

<!-- grade_id: cell-ebed9b4492e23bd9, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-ed2833098327624d, points: 1, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```