Please read the following instructions thoroughly. Neglecting to do so may result in missed points.

### Preamble
**Reminder**: Homeworks are due by 11:59PM ET on Sundays.

Before you turn this problem set in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

### Naming conventions
Be sure the filename of your notebook is in the following form:

    <uni>_<assignment>_<details [optional]>.<extension>
    
For example:

    lr3086_hw01.ipynb
    lr3086_hw01_complete.ipynb
    LR3086_HW01.ipynb
    
To rename a notebook, in the menubar, select File$\rightarrow$Rename. The extension for notebook files, `.ipynb`, will already be appended to the filename, but will be hidden from view within the notebook.
    
This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `lr3086_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

### What Format To Submit In

Most homeworks are in Jupyter notebooks. Once you've finished your homework, unless specified otherwise, please download your work as an `.ipynb` file to your local machine, then upload it to Courseworks when complete (in the menubar, select File$\rightarrow$Download as$\rightarrow$Notebook).

**Failure to submit a Jupyter notebook will result in a grade of zero for the assignment.**

### Grading

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors.

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem.

Empty un-editable cells in an assignment are there for a reason. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

### Getting Help

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Ed Discussions or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Ed Discussions and not via email. Answers given via Ed Discussions will help you as well as other students. Thus, emails will always have a lower priority for response than Ed Discussions questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in Courseworks faster.

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

---

# Homework 6: Regular Expressions

Total questions: 5 <br/>
Total points: 5

## How to complete this homework

In each answer, you'll be asked to define a single variable called `pattern` which should match the given prompt. **Note**: we're _not_ defining functions, just the variable `pattern`.

As an example, if asked for a regular expression that matches the letter a contained anywhere within a string, you should write:

```py
pattern = "a"
```

as the sole contents of the cell. Recall that if you wish to use character classes like `\w`, you likely will want to use a raw string like `r"foo\w\wbar"` so that you don't need to escape the slashes twice.

## Question 1

Write a regular expression pattern which matches the substring "cat" anywhere within a string. Your pattern should also successfully match the word "caterpillar" as well as "scathing".

[1 point]

<!-- grade_id: cell-4215ecfba4aa763f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
pattern= r"cat"

```

<!-- grade_id: cell-b92229c1be2c43b8, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
import re
assert re.search(pattern, "cat")
### END TESTS
```

<!-- grade_id: cell-a49dc667e28e368f, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert re.search(pattern, "caterpillar")
### END TESTS
```

<!-- grade_id: cell-0b6fcc066144f157, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert re.search(pattern, "scathing")
assert re.search(pattern, "The cat jumped on the bed.")
### END TESTS
```

<!-- grade_id: cell-ddcc96e15fcae904, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert not re.search(pattern, "dog")
assert not re.search(pattern, "aca")
assert not re.search(pattern, "ca")
assert not re.search(pattern, "at")
### END TESTS
```

## Question 2

Write a regular expression which matches the substring "cat" at the beginning of a string, and not elsewhere. Your pattern should also successfully match the word "caterpillar" at the beginning of a string.

[1 point]

<!-- grade_id: cell-f07df74b184bf808, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
pattern= r"^cat"
```

<!-- grade_id: cell-9b8b67c776e84aef, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
import re
assert re.search(pattern, "cat")
assert re.search(pattern, "caterpillar")
### END TESTS
```

<!-- grade_id: cell-2648975c8550ad70, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert not re.search(pattern, "scathing")
assert not re.search(pattern, "The cat jumped on the bed.")
### END TESTS
```

## Question 3

In class, we mostly used the `re.search` function. The `re` module has an additional function called `re.findall` which, rather than finding only one place in a string that matches a given pattern, finds all of the places which match.

Given a string like:

```py
example = """
2010-04-29 1984 ABC-123!
Feb 18th, 2018. 
888-27-8949 2048-07-07
"""
```

write a pattern which matches dates within the string which look like YYYY-MM-DD.

Your answer should again be a single variable `pattern`, but such that `re.findall(pattern, example)` returns `["2010-04-29", "2048-07-07"]` which are the two matching parts of the string.

[1 point]

<!-- grade_id: cell-6572b7193413bba5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
pattern=r"\d{4}-\d{2}-\d{2}"



```

<!-- grade_id: cell-56c8d107244c0f73, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
string = """
2010-04-29 1984 ABC-123!
Feb 18th, 2018. 
888-27-8949 2048-07-07
"""
expected = ['2010-04-29', '2048-07-07']
actual = re.findall(pattern, string)
assert expected == actual 
### END TESTS
```

## Question 4

A building manager has textual logs of problems occuring in apartment units it manages.

An individual message looks like:

    Problem reported in apartment 3F with plumbing.

or

    Problem in apartment 12a with sink.

or

    Noise reported from apartment 8g.

or

    An apartment owner reports that apartment 4A has been broken into.

Use a pattern with a capture group to extract the apartment information from the string -- include both the floor and unit, so for the first example, your pattern should have `re.search(pattern).group(1)` produce `"3F"`; and for the last one, `"4A".`

If the message does not include an apartment unit, you should not match anything.

[1 point]

<!-- grade_id: cell-8f3ec6bfcd6dcf77, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
pattern=r"((\d{1,})([a-z]|[A-Z]))"

```

<!-- grade_id: cell-e0a4a3088fe5c5b5, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert "3F" == re.search(pattern, "Problem reported in apartment 3F with plumbing.").group(1)
assert "12a" == re.search(pattern, "Problem in apartment 12a with sink.").group(1)
### END TESTS
```

<!-- grade_id: cell-28eeacdcb1be3574, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert "4A" == re.search(pattern, "An apartment owner reports that apartment 4A has been broken into.").group(1)
### END TESTS
```

<!-- grade_id: cell-0a079abd72adb82c, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert re.search(pattern, "No incidents.") is None
### END TESTS
```

## Question 5

Write a pattern which matches strings containing odd numbers anywhere within it, but not even ones.

[1 point]

<!-- grade_id: cell-45fe83c338be660d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
pattern=r"\d*[13579]"
```

<!-- grade_id: cell-cab2c19222bddaf2, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert re.search(pattern, "The number 1 matches")
### END TESTS
```

<!-- grade_id: cell-3b2acf848a2387ca, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert re.search(pattern, "The number 9 matches")
### END TESTS
```

<!-- grade_id: cell-e99b7e5763951347, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert not re.search(pattern, "2 does not match")
### END TESTS
```

<!-- grade_id: cell-31afb06827a674af, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
assert re.search(pattern, "5 does match")
### END TESTS
```

```python

```