Please read the following instructions thoroughly. Neglecting to do so may result in missed points.

### Preamble
**Reminder**: Homeworks are due by 11:59PM ET on Sundays.

Before you turn this problem set in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

### Naming conventions
Be sure the filename of your notebook is in the following form:

    <uni>_<assignment>_<details [optional]>.<extension>
    
For example:

    lr3086_hw01.ipynb
    lr3086_hw01_complete.ipynb
    LR3086_HW01.ipynb
    
To rename a notebook, in the menubar, select File$\rightarrow$Rename. The extension for notebook files, `.ipynb`, will already be appended to the filename, but will be hidden from view within the notebook.
    
This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `lr3086_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

### What Format To Submit In

Most homeworks are in Jupyter notebooks. Once you've finished your homework, unless specified otherwise, please download your work as an `.ipynb` file to your local machine, then upload it to Courseworks when complete (in the menubar, select File$\rightarrow$Download as$\rightarrow$Notebook).

**Failure to submit a Jupyter notebook will result in a grade of zero for the assignment.**

### Grading

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors.

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem.

Empty un-editable cells in an assignment are there for a reason. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

### Getting Help

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Ed Discussions or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Ed Discussions and not via email. Answers given via Ed Discussions will help you as well as other students. Thus, emails will always have a lower priority for response than Ed Discussions questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in Courseworks faster.

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

---

# Homework 10: SQLAlchemy

Total questions: 5<br/>
Total points: 8

### FYI

In class, we used paths on your local machine outside of the notebook to persist data, but here, we'll use a special argument that `'sqlite:///:memory:'` in SQLAlchemy has which stores the database in-memory -- in otherwords, the database will live only so long as you're running the notebook. Afterwards, it will be thrown away. Usually this argument is useful when experimenting, given that as we discussed, it is persistence that often makes databases useful. We use it here however simply to avoid needing to upload multiple files to Courseworks -- everything you need is here.

If for any reason you need to discard the contents of your database to start over, you may therefore do so by simply restarting the "kernel" of this notebook, which you can do by clicking the restart button in the toolbar (it's 2 buttons to the right of the "Run" button) or in the Kernel menu.

## Question 1

Using `SQLAlchemy`'s ORM (object relational mapping) layer, we will be creating a new database with some new tables.

For the first table, create a table named `students` by defining a `Student` class. A `Student` should have the following:

* a `uni` (`String`)
* a `fullname` (`String`)
* a `nickname` (`String`)
* a planned `graduation_date` (`Date`)
* a tuition `balance` (`Float`)

as well as an `id` as a primary key.

Be sure to define the `__tablename__` attribute within the `Student` class.

[2.5 points]

```python
from sqlalchemy import create_engine

#engine = create_engine(f"sqlite:///try.db", echo=True)
```

<!-- grade_id: cell-d15613ee427db50a, points: 0, grade: false, solution: true -->
```python
# <-- DO NOT DELETE
from sqlalchemy import Column, Date, Integer, Float, String
from sqlalchemy.ext.declarative import declarative_base

Base = declarative_base()
# -->

class Student(Base):
    # YOUR CODE HERE
   # raise NotImplementedError()
    __tablename__="students"
    
    id=Column(Integer, primary_key=True)
    uni=Column(String)
    fullname=Column(String)
    nickname=Column(String)
    graduation_date=Column(Date)
    balance=Column(Float)
    
#Base.metadata.create_all(engine, checkfirst=True)
```

<!-- grade_id: cell-52086feff76a74c0, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-cb9810b62b85cac3, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-588fdf97e23cd5d1, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-e2642590c29e05d8, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-39479c17323060c6, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

## Question 2

Next, let's define a `classes` table with a `Class` class. A `Class` should have the following columns:

* a `name` (`String`) of the class
* a `days` (`String`) of week representing the schedule (e.g. `M` for Monday schedule, `MW` for Monday + Wednesday schedule, `TR` for Tuesday + Thursday schedule, etc)
* a `professor` (`String`) name

as well as an `id` as a primary key.

Again, be sure to define the `__tablename__` attribute within the `Class` class.

[1 point]

<!-- grade_id: cell-90272762fb8c542f, points: 0, grade: false, solution: true -->
```python
class Class(Base):
    # YOUR CODE HERE
    #raise NotImplementedError()
    __tablename__="classes"
    
    id=Column(Integer, primary_key=True)
    name=Column(String)
    days=Column(String)
    professor=Column(String)
    
#Base.metadata.create_all(engine, checkfirst=True)
```

<!-- grade_id: cell-5f4a0472c7c13659, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-0cda02dbf8d60ae5, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

## Question 3

It makes sense that a class has students enrolled in it. And that students can be enrolled in one or more classes. We consider this a [many-to-many relationship](https://docs.sqlalchemy.org/en/14/orm/basic_relationships.html#many-to-many). (It may also be helpful to read [this part](https://docs.sqlalchemy.org/en/14/orm/tutorial.html#building-a-many-to-many-relationship) of the SQLAlchemy tutorial.)

Defined for you is an association table, `student_classes`. You are to re-define both the `Student` and `Class` classes:

Re-define your `Student` class (using the same table name and columns already defined) to include a `classes` attribute using SQLAlchemy's [`relationship()` method](https://docs.sqlalchemy.org/en/14/orm/relationship_api.html#sqlalchemy.orm.relationship). The relationship should be called with 3 arguments: one for the Python class that we're creating a relationship to; one for the [`back_populates`](https://docs.sqlalchemy.org/en/14/orm/relationship_api.html#sqlalchemy.orm.relationship.params.back_populates) argument; and one for [`secondary`](https://docs.sqlalchemy.org/en/14/orm/relationship_api.html#sqlalchemy.orm.relationship.params.secondary) to map to the association table.

Do the same with your `Class` class with including a `students` attrubte using the `relationship()` method.


[1.5 points]

<!-- grade_id: cell-cdeadd2349087623, points: 0, grade: false, solution: true -->
```python
# <-- DO NOT DELETE
from sqlalchemy.schema import ForeignKey
from sqlalchemy import Table
from sqlalchemy.orm import relationship


Base = declarative_base()
student_class = Table("student_classes", Base.metadata,
    Column("student_id", ForeignKey("students.id"), primary_key=True),
    Column("class_id", ForeignKey("classes.id"), primary_key=True),
                      
)
# -->


class Student(Base):
    # YOUR CODE HERE
   # raise NotImplementedError()
    __tablename__="students"
    
    id=Column(Integer, primary_key=True)
    uni=Column(String)
    fullname=Column(String)
    nickname=Column(String)
    graduation_date=Column(Date)
    balance=Column(Float)
    classes=relationship('Class', secondary= student_class, back_populates='students')
    

class Class(Base):
    # YOUR CODE HERE
   # raise NotImplementedError()
    __tablename__="classes"
    
    id=Column(Integer, primary_key=True)
    name=Column(String)
    days=Column(String)
    professor=Column(String)
    students=relationship('Student', secondary=student_class, back_populates='classes')
    
#Base.metadata.create_all(engine)
```

<!-- grade_id: cell-168a4fb5f14594d9, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-78928eea2439957b, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

<!-- grade_id: cell-f97eb07cdc365d35, points: 0.5, grade: true, solution: false -->
```python
# CELL INTENTIONALLY LEFT BLANK - DO NOT ALTER OR DELETE
```

## Question 4

Let's now add some data into our two tables.

With the data provided below, create 3 students using your `Student` class, create 4 classes using your `Class` class, and make sure the students are "enrolled" in the listed classes. You'll then use `session.add_all` and `session.commit` to add the 3 students and 2 classes to your database.

**Students:**

| UNI    | Fullname          | Nickname | Graduation Date   | Balance | Classes   |
|--------|-------------------|----------|-------------------|---------|-----------|
| ab1234 | Elizabeth Rose    | Liz      | May 1, 2023       | 0       | Stochastic Models |
| cd5678 | Jon-Paul Phillips | JP       | December 15, 2022 | 526.50  | Stochastic Models, Data Analytics & Machine Learning |
| ef0987 | Idris Sanders     |          | May 1, 2024       | 392.12  | Data Analytics & Machine Learning, Foundations of Data Science |

**Classes:**

| Name                              | Days | Professor       |
|-----------------------------------|------|-----------------|
| Stochastic Models                 | TR   | Marcus Brown    |
| Data Analytics & Machine Learning | MW   | Grace Robbinson |
| Foundations of Data Science       | T    | Shannon Wells   |
| Reinforcement Learning            | WF   | Joe Greene      |


Name your variables like so: `student1`, `student2`, `class1`, etc.

[2 points]

<!-- grade_id: cell-d3fc829da20488ee, points: 0, grade: false, solution: true -->
```python
# <-- DO NOT DELETE
import datetime 
from sqlalchemy.orm import sessionmaker
from sqlalchemy import create_engine
engine = create_engine("sqlite:///:memory:")
Base.metadata.create_all(engine)
Session = sessionmaker(bind=engine)
session = Session()
# -->

# YOUR CODE HERE
#raise NotImplementedError()


class1=Class(name='Stochastic Models' , days= 'TR' , professor= 'Marcus Brown' )
class2=Class(name= 'Data Analytics & Machine Learning', days='MW'  , professor= 'Grace Robbinson' )
class3=Class(name= 'Foundations of Data Science', days= 'T' , professor= 'Shannon Wells' )
class4=Class(name='Reinforcement Learning' , days= 'WF' , professor= 'Joe Greene' )


student1=Student(uni='ab1234', fullname='Elizabeth Rose', nickname='Liz', graduation_date=datetime.datetime(2023,5,1),
                 balance=0)
student2=Student(uni='cd5678', fullname='Jon-Paul Phillips', nickname='JP', graduation_date=datetime.datetime(2022,12,15),
                 balance=526.50)
student3=Student(uni='ef0987', fullname='Idris Sanders', 
                 graduation_date=datetime.datetime(2024,5,1), balance=392.12)

student1.classes.append(class1)
student2.classes.append(class1)
student2.classes.append(class2)
student3.classes.append(class2)
student3.classes.append(class3)


# <-- DO NOT DELETE
session.add_all([
    student1, student2, student3,
    class1, class2, class3, class4
])
session.commit()
# -->
```

<!-- grade_id: cell-eae2078c2f2c0707, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
results = session.query(Student).filter_by(uni="ab1234").all()

assert results[0] == student1
assert results[0].classes == [class1]
### END TESTS
```

<!-- grade_id: cell-a8a5e0663511837f, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
results = session.query(Student).filter_by(uni="cd5678").all()

assert results[0] == student2
assert results[0].classes == [class1, class2]
### END TESTS
```

<!-- grade_id: cell-6eecf06da65163e5, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
results = session.query(Student).filter_by(uni="ef0987").all()

assert results[0] == student3
assert results[0].classes == [class2, class3]
### END TESTS
```

<!-- grade_id: cell-d9d0b682e37925ef, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
results = session.query(Class).filter_by(students=None).all()
assert results == [class4]
### END TESTS
```

## Question 5

Recall the types of joins we reviewed on Monday's lecture. Define a query against your database using SQLAlchemy's ORM with the `session.query` method that returns all classes (just the class `name`) and the number of enrolled students. Include classes that do not have any students enrolled. 

Assign your query to a variable called `result`.

Be sure that the query is executed, so that the `result` is a `list` of `tuple`s containing two items each: the name of the class, and number of students enrolled. Order does not matter.

You may need to use [SQLAlchemy's `func` module](https://docs.sqlalchemy.org/en/14/core/functions.html) so it's been imported for you; but feel free to import other SQLAlchemy modules as needed.

[1 point]

<!-- grade_id: cell-1f7ff4ae0f568ffa, points: 0, grade: false, solution: true -->
```python
from sqlalchemy import func

# YOUR CODE HERE
#raise NotImplementedError()
result=session.query(Class.name, func.count(student_class.columns[1])).outerjoin(student_class).group_by(student_class.columns[1]).all()




```

<!-- grade_id: cell-029caf9592584241, points: 1, grade: true, solution: false -->
```python
### BEGIN TESTS
expected = sorted([
    ('Reinforcement Learning', 0),
    ('Stochastic Models', 1),
    ('Data Analytics & Machine Learning', 2),
    ('Stochastic Models', 2)
])
assert expected == sorted(result)
### END TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [21], in <cell line: 8>()
      1 ### BEGIN TESTS
      2 expected = sorted([
      3     ('Reinforcement Learning', 0),
      4     ('Stochastic Models', 1),
      5     ('Data Analytics & Machine Learning', 2),
      6     ('Stochastic Models', 2)
      7 ])
----> 8 assert expected == sorted(result)

AssertionError:
```